<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 86  -->
  <question type="coderunner">
    <name>
      <text>BUILTIN_PROTOTYPE_blockly_python3_w_input</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Built-in prototypes are documented in the language strings.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>blockly</coderunnertype>
    <prototypetype>1</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>33.3, 66.7, ...</penaltyregime>
    <precheck>0</precheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[<xml xmlns="https://developers.google.com/blockly/xml"><code xmlns=""/></xml>]]></answerpreload>
    <globalextra></globalextra>
    <useace>1</useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[import xml.etree.ElementTree as ET

__saved_input__ = input
def input(prompt=''):
    s = __saved_input__(prompt)
    return s

answer = """
{{ STUDENT_ANSWER | e('py') }}
"""

code = ET.fromstring(answer).find("code").text

if code:
    exec(code)

SEPARATOR = "#<ab@17943918#@>#"

{% for TEST in TESTCASES %}
{{ TEST.testcode }}
{% if not loop.last %}
print(SEPARATOR)
{% endif %}
{% endfor %}]]></template>
    <iscombinatortemplate>1</iscombinatortemplate>
    <allowmultiplestdins>0</allowmultiplestdins>
    <answer><![CDATA[<xml xmlns="https://developers.google.com/blockly/xml"><code xmlns=""/></xml>]]></answer>
    <validateonsave>0</validateonsave>
    <testsplitterre><![CDATA[|#<ab@17943918#@>#\n|ms]]></testsplitterre>
    <language>python3</language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader>EqualityGrader</grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams><![CDATA[{"locale": "en",
 "toolbox": "<xml xmlns=\"https:\/\/developers.google.com\/blockly\/xml\" id=\"toolbox\" style=\"display: none\">\r\n  <category name=\"%{BKY_CATEGORY_NAME_LOGIC}\" colour=\"#5b80a5\">\r\n    <block type=\"controls_if\"><\/block>\r\n    <block type=\"logic_compare\">\r\n      <field name=\"OP\">EQ<\/field>\r\n    <\/block>\r\n    <block type=\"logic_operation\">\r\n      <field name=\"OP\">AND<\/field>\r\n    <\/block>\r\n    <block type=\"logic_negate\"><\/block>\r\n    <block type=\"logic_boolean\">\r\n      <field name=\"BOOL\">TRUE<\/field>\r\n    <\/block>\r\n    <block type=\"logic_null\"><\/block>\r\n    <block type=\"logic_ternary\"><\/block>\r\n  <\/category>\r\n  <category name=\"%{BKY_CATEGORY_NAME_LOOPS}\" colour=\"#5ba55b\">\r\n    <block type=\"controls_repeat_ext\">\r\n      <value name=\"TIMES\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">10<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"controls_whileUntil\">\r\n      <field name=\"MODE\">WHILE<\/field>\r\n    <\/block>\r\n    <block type=\"controls_for\">\r\n      <field name=\"VAR\" id=\"e%f!fa?5S+!M?9haF%eQ\">i<\/field>\r\n      <value name=\"FROM\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">1<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n      <value name=\"TO\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">10<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n      <value name=\"BY\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">1<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"controls_forEach\">\r\n      <field name=\"VAR\" id=\"4L#?uuDWH,f2n5F~:;%E\">j<\/field>\r\n    <\/block>\r\n    <block type=\"controls_flow_statements\">\r\n      <field name=\"FLOW\">BREAK<\/field>\r\n    <\/block>\r\n  <\/category>\r\n  <category name=\"%{BKY_CATEGORY_NAME_MATH}\" colour=\"#5b67a5\">\r\n    <block type=\"math_number\">\r\n      <field name=\"NUM\">0<\/field>\r\n    <\/block>\r\n    <block type=\"math_arithmetic\">\r\n      <field name=\"OP\">ADD<\/field>\r\n      <value name=\"A\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">1<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n      <value name=\"B\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">1<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"math_single\">\r\n      <field name=\"OP\">ROOT<\/field>\r\n      <value name=\"NUM\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">9<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"math_trig\">\r\n      <field name=\"OP\">SIN<\/field>\r\n      <value name=\"NUM\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">45<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"math_constant\">\r\n      <field name=\"CONSTANT\">PI<\/field>\r\n    <\/block>\r\n    <block type=\"math_number_property\">\r\n      <mutation divisor_input=\"false\"><\/mutation>\r\n      <field name=\"PROPERTY\">EVEN<\/field>\r\n      <value name=\"NUMBER_TO_CHECK\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">0<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"math_round\">\r\n      <field name=\"OP\">ROUND<\/field>\r\n      <value name=\"NUM\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">3.1<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"math_on_list\">\r\n      <mutation op=\"SUM\"><\/mutation>\r\n      <field name=\"OP\">SUM<\/field>\r\n    <\/block>\r\n    <block type=\"math_modulo\">\r\n      <value name=\"DIVIDEND\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">64<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n      <value name=\"DIVISOR\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">10<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"math_constrain\">\r\n      <value name=\"VALUE\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">50<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n      <value name=\"LOW\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">1<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n      <value name=\"HIGH\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">100<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"math_random_int\">\r\n      <value name=\"FROM\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">1<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n      <value name=\"TO\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">100<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"math_random_float\"><\/block>\r\n  <\/category>\r\n  <category name=\"%{BKY_CATEGORY_NAME_TEXT}\" colour=\"#5ba58c\">\r\n    <block type=\"text\">\r\n      <field name=\"TEXT\"><\/field>\r\n    <\/block>\r\n    <block type=\"text_join\">\r\n      <mutation items=\"2\"><\/mutation>\r\n    <\/block>\r\n    <block type=\"text_append\">\r\n      <field name=\"VAR\" id=\"D\/TqD@[_xou$^4Ry=-:T\">item<\/field>\r\n      <value name=\"TEXT\">\r\n        <shadow type=\"text\">\r\n          <field name=\"TEXT\"><\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"text_length\">\r\n      <value name=\"VALUE\">\r\n        <shadow type=\"text\">\r\n          <field name=\"TEXT\">abc<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"text_isEmpty\">\r\n      <value name=\"VALUE\">\r\n        <shadow type=\"text\">\r\n          <field name=\"TEXT\"><\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"text_indexOf\">\r\n      <field name=\"END\">FIRST<\/field>\r\n      <value name=\"VALUE\">\r\n        <block type=\"variables_get\">\r\n          <field name=\"VAR\" id=\"ch[oxT*Mv~`n^#F8G7h3\">text<\/field>\r\n        <\/block>\r\n      <\/value>\r\n      <value name=\"FIND\">\r\n        <shadow type=\"text\">\r\n          <field name=\"TEXT\">abc<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"text_charAt\">\r\n      <mutation at=\"true\"><\/mutation>\r\n      <field name=\"WHERE\">FROM_START<\/field>\r\n      <value name=\"VALUE\">\r\n        <block type=\"variables_get\">\r\n          <field name=\"VAR\" id=\"ch[oxT*Mv~`n^#F8G7h3\">text<\/field>\r\n        <\/block>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"text_getSubstring\">\r\n      <mutation at1=\"true\" at2=\"true\"><\/mutation>\r\n      <field name=\"WHERE1\">FROM_START<\/field>\r\n      <field name=\"WHERE2\">FROM_START<\/field>\r\n      <value name=\"STRING\">\r\n        <block type=\"variables_get\">\r\n          <field name=\"VAR\" id=\"ch[oxT*Mv~`n^#F8G7h3\">text<\/field>\r\n        <\/block>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"text_changeCase\">\r\n      <field name=\"CASE\">UPPERCASE<\/field>\r\n      <value name=\"TEXT\">\r\n        <shadow type=\"text\">\r\n          <field name=\"TEXT\">abc<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"text_trim\">\r\n      <field name=\"MODE\">BOTH<\/field>\r\n      <value name=\"TEXT\">\r\n        <shadow type=\"text\">\r\n          <field name=\"TEXT\">abc<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"text_print\">\r\n      <value name=\"TEXT\">\r\n        <shadow type=\"text\">\r\n          <field name=\"TEXT\">abc<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"text_prompt_ext\">\r\n      <value name=\"TEXT\">\r\n        <shadow type=\"text\">\r\n          <field name=\"TEXT\"><\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n  <\/category>\r\n  <category name=\"%{BKY_CATEGORY_NAME_LISTS}\" colour=\"#745ba5\">\r\n    <block type=\"lists_create_with\">\r\n      <mutation items=\"0\"><\/mutation>\r\n    <\/block>\r\n    <block type=\"lists_create_with\">\r\n      <mutation items=\"3\"><\/mutation>\r\n    <\/block>\r\n    <block type=\"lists_repeat\">\r\n      <value name=\"NUM\">\r\n        <shadow type=\"math_number\">\r\n          <field name=\"NUM\">5<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"lists_length\"><\/block>\r\n    <block type=\"lists_isEmpty\"><\/block>\r\n    <block type=\"lists_indexOf\">\r\n      <field name=\"END\">FIRST<\/field>\r\n      <value name=\"VALUE\">\r\n        <block type=\"variables_get\">\r\n          <field name=\"VAR\" id=\"^A2#7{?Q`;n4`M-*1J4E\">list<\/field>\r\n        <\/block>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"lists_getIndex\">\r\n      <mutation statement=\"false\" at=\"true\"><\/mutation>\r\n      <field name=\"MODE\">GET<\/field>\r\n      <field name=\"WHERE\">FROM_START<\/field>\r\n      <value name=\"VALUE\">\r\n        <block type=\"variables_get\">\r\n          <field name=\"VAR\" id=\"^A2#7{?Q`;n4`M-*1J4E\">list<\/field>\r\n        <\/block>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"lists_setIndex\">\r\n      <mutation at=\"true\"><\/mutation>\r\n      <field name=\"MODE\">SET<\/field>\r\n      <field name=\"WHERE\">FROM_START<\/field>\r\n      <value name=\"LIST\">\r\n        <block type=\"variables_get\">\r\n          <field name=\"VAR\" id=\"^A2#7{?Q`;n4`M-*1J4E\">list<\/field>\r\n        <\/block>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"lists_getSublist\">\r\n      <mutation at1=\"true\" at2=\"true\"><\/mutation>\r\n      <field name=\"WHERE1\">FROM_START<\/field>\r\n      <field name=\"WHERE2\">FROM_START<\/field>\r\n      <value name=\"LIST\">\r\n        <block type=\"variables_get\">\r\n          <field name=\"VAR\" id=\"^A2#7{?Q`;n4`M-*1J4E\">list<\/field>\r\n        <\/block>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"lists_split\">\r\n      <mutation mode=\"SPLIT\"><\/mutation>\r\n      <field name=\"MODE\">SPLIT<\/field>\r\n      <value name=\"DELIM\">\r\n        <shadow type=\"text\">\r\n          <field name=\"TEXT\">,<\/field>\r\n        <\/shadow>\r\n      <\/value>\r\n    <\/block>\r\n    <block type=\"lists_sort\">\r\n      <field name=\"TYPE\">NUMERIC<\/field>\r\n      <field name=\"DIRECTION\">1<\/field>\r\n    <\/block>\r\n  <\/category>\r\n  <sep><\/sep>\r\n  <category name=\"%{BKY_CATEGORY_NAME_VARIABLES}\" colour=\"#a55b80\" custom=\"VARIABLE\"><\/category>\r\n  <category name=\"%{BKY_CATEGORY_NAME_FUNCTIONS}\" colour=\"#995ba5\" custom=\"PROCEDURE\"><\/category>\r\n<\/xml>"}]]></templateparams>
    <hoisttemplateparams>0</hoisttemplateparams>
    <twigall>0</twigall>
    <uiplugin>blockly</uiplugin>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>0</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>0</displayfeedback>
    <testcases>
    </testcases>
  </question>

</quiz>